**Wichtige Hinweise**

- Aktivieren Sie die entsprechende Conda-Umgebung bevor Sie beginnen.

- Geben Sie als `NAME` ihr HdM-Kürzel an.

- Ändern Sie **nicht** den Namen der Datei und löschen Sie keine Zellen.

- Bearbeiten Sie alle Zellen mit dem Hinweis <font color='green'> \# YOUR CODE HERE </font>

- Die Funktion **NotImplementedError()** soll die Abgabe von leeren Zellen verhindern. Löschen Sie die Funktion, sobald Sie in einer dieser Zellen arbeiten.

- Stellen Sie sicher, dass alles wie erwartet läuft, bevor Sie die Prüfung abgeben: Starten Sie den Kernel neu und führen Sie alle Zellen aus: wählen Sie "Restart" und dann "Run All"

```python
NAME = "mn086"
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Netflix user engagement analyse

## Setup

```python
# Wir importieren die Bibliothek pandas
import pandas as pd
```

---
## Daten

### Daten importieren (2 Punkte)

- Wir haben Pandas mit der Bezeichnung `pd` importiert.

- Die Pandas-Funktion zum Einlesen von CSV-Dateien: `.read_csv`

- CSV-Daten mit Pandas importieren und mit der Bezeichnung `df` speichern.

- Pfad zu den Daten in GitHub: 'https://raw.githubusercontent.com/kirenz/datasets/master/netflix.csv'


- Hinweis:

```python
___ = ___.___(___)
```

<!-- grade_id: import, points: 0, grade: false, solution: true -->
```python
df=pd.read_csv('https://raw.githubusercontent.com/kirenz/datasets/master/netflix.csv')
```

<!-- grade_id: correct_import, points: 2, grade: true, solution: false -->
```python
"""Check that df returns the correct output"""
assert len(df) == 2000
assert df.columns.tolist() == ['group', 'outcome']
```

---
## Analyse

### Kreuztabelle (4 Punkte)

Schreiben Sie den Code, der eine Kreuztabelle erzeugt erzeugt:

- Speichern Sie die Tabelle unter der Bezeichnung: `cross_table`
- Die Funktion in Pandas zur Erzeugung von Kreuztabellen: `.crosstab(ZEILEN, SPALTEN)`
- Als Zeilen diese Variable nutzen: `group`
- Als Spalten diese Variable nutzen: `outcome`
- Geben Sie die Spalten- und Zeilensummen aus (`margins=True`)
- Geben Sie die Werte als Prozentwerte aus (`normalize=True`)
- Die Variable `group` soll in der Tabelle "Gruppe" genannt werden (`rownames`)
- Die Variable `outcome` soll in der Tabelle "Ergebnis" genannt werden (`colnames`)
- Formatierung der Prozent-Werte mit Multiplikation: bspw. sollte ein Wert von 0.1775 als 17.75 angezeigt werden
- Geben Sie die Tabelle aus

- Hinweis:

```python

___ = __.__(df['___'], df['___'],
            ___=True, # Spalten- und Zeilensumme
            ___=True, # Prozentwerte
            rownames=['___'],  # Umbenennung der Zeilen
            colnames=['___'] # Umbenennung der Spalten
            ) * ___ # Multiplikation mit 100 für eine bessere Darstellung der Prozentwerte

___

```

<!-- grade_id: cross_table, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: correct_cross_table, points: 4, grade: true, solution: false -->
```python
"""Check that it returns the correct output"""
assert cross_table['Engagement'].B == 21.8
```